# Coursework
# SPCE0038: Machine Learning with Big-Data

---

## Overview

This coursework is provided as a Jupyter notebook, which you will need to complete.  

Throughout the notebook you will need to complete code, analytic exercises (if equations are required please typeset your solutions using latex in the markdown cell provided) and descriptive answers. Much of the grading of the coursework will be performed automatically, so it is critical you name your variables as requested.

Before you turn this coursework in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says "YOUR ANSWER HERE" or `YOUR CODE HERE` and remove remove the `raise NotImplementedError()` exceptions that are thrown before you have added your answers. Do not add and remove cells but rather provide your answers in the spaces given.

Please also:
- Make sure your notebook executes without errors.
- Do not overwrite or rename any existing variables.
- Do not install code or packages in the notebooks.

Please rename the notebook filename to include your name in the filename.  And please also add your name below:

```python
NAME = ""
```

You will be able to run some basic tests in the notebook to check the basic operation of your code is as expected.  Although do not assume your responses are complete or fully correct just because the basic tests pass.

Once you have renamed the notebook file and completed the exercises, please upload the notebook to Moodle.


---

## AstroML

The data used is this coursework is obtained using [AstroML](http://www.astroml.org), a python package for machine learning for astronomy.  Although we take data from AstroML, this coursework is not based on standard AstroML examples.  So you will *not* find the solutions in AstroML examples!

## SDSS

The data obtained through AstroML was observed by the [Sloan Digital Sky Survey](https://www.sdss.org/) (SDSS), which began observations in 2000.  SDSS data have lead to many scientific advances and the experiment is widely seen as one of the most successful surveys in the history of astronomy.

---

## Dependencies

- Standard course dependencies (e.g. numpy, scikit-learn, etc.)
- [AstoML](http://www.astroml.org)
- [AstroPy](http://www.astropy.org/)

---

```python
import numpy as np
from matplotlib import pyplot as plt
```

```python
def check_var_defined(var):
    try:
        exec(var)
    except NameError:
        raise NameError(var + " not defined.")
    else:
        print(var + " defined.")
```

## Part 1: Regression

In these exercises we will consider the regression problem of the astonomical distance modulus vs redshift relationship.

In astronomy, the [distance modulus](https://en.wikipedia.org/wiki/Distance_modulus) specifies the difference between the apparent and absolute magnitudes of an astronomnical object.  It provides a way of expressing astrophysical distances. 

Astronomical [redshift](https://en.wikipedia.org/wiki/Redshift) specifies the shift in wavelength that astronomical objects undergo due to the expansion of the Universe.  Due to Hubble's Law, more distance objects experience a greater redshift.


```python
from astroML.datasets import generate_mu_z
```

```python
# Load data
m = 150
z_sample, mu_sample, dmu = generate_mu_z(m, random_state=3)
```

*Plot the distance modulus ($\mu$) vs redhift ($z$), including error bars.*

<!-- grade_id: cell-ae7f16835ab51af3, points: 1, grade: true, solution: true -->
```python
# Plot data
def plot_dist_mod():
    # YOUR CODE HERE
    raise NotImplementedError()
    plt.xlabel('$z$')
    plt.ylabel('$\mu$')
    plt.title('Distance modulus vs redshift')
    plt.ylim(36, 50)
    plt.xlim(0, 1.5)
plot_dist_mod()
```

Recall the normal equations for linear regression follow by analytically minimising the cost function: 

$$\min_\theta\ C(\theta) = \min_\theta \ (X \theta - y)^{\rm T}(X \theta - y).$$

Show analytically that the solution is given by 

$$ \hat{\theta} = \left( X^{\rm T} X \right)^{-1} X^{\rm T} y. $$

[Matrix calculus identities](https://en.wikipedia.org/wiki/Matrix_calculus) may be useful (note that we use the denominator layout convention).

*Expand the cost function and drop terms that do not depend on $\theta$ (use latex mathematics expressions):*

<!-- grade_id: cell-a114fd93ba74dac5, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

*Set the derivative to zero and solve for $\theta$ (use latex mathematics expressions):*

<!-- grade_id: cell-0f0b521f765f4826, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

*Solve for $\theta$ by numerically implementing the analytic solution given above.*

<!-- grade_id: cell-d71c2644693323b2, points: 0, grade: false, solution: true -->
```python
def compute_theta_lin_reg(X, y):
    # YOUR CODE HERE
    raise NotImplementedError()
    return theta
```

<!-- grade_id: cell-f024710582a10726, points: 1, grade: true, solution: false -->
```python
assert compute_theta_lin_reg(z_sample, mu_sample).shape == (2,)
theta = compute_theta_lin_reg(z_sample, mu_sample)
(theta_c, theta_m) = theta
print("Linear regression parameters recovered analytically: intercept={0:.4f}, slope={1:.4f}".format(theta_c, theta_m))
```

<!-- grade_id: cell-52c5a2bcccb2010d, points: 1, grade: true, solution: false -->
```python
check_var_defined('theta_c')
check_var_defined('theta_m')
```

*Write a method to make a prediction for a given redshift.*

<!-- grade_id: cell-f5341cc7da485877, points: 0, grade: false, solution: true -->
```python
def predict_lin_reg(theta, x):
    # YOUR CODE HERE
    raise NotImplementedError()
    return y
```

*Predict the distance modulus for a range of redshift values between 0.01 and 1.5 and plot the predicted curve overlayed on data (make a new plot; do not revise the plot above).  Call the variable used to store the predictions for your polynomial model `mu_pred_lin`.*

<!-- grade_id: cell-fc6aa680d6528ab1, points: 0, grade: false, solution: true -->
```python
z = np.linspace(0.01, 1.5, 1000)
plot_dist_mod()
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f8e904006f960080, points: 2, grade: true, solution: false -->
```python
check_var_defined('mu_pred_lin')
assert mu_pred_lin.shape == (len(z),), "Make sure the shape of your predictions is correct"
```

*Solve for the parameters $\theta$ using Scikit-Learn.*

<!-- grade_id: cell-21b18d58a127b96c, points: 0, grade: false, solution: true -->
```python
from sklearn.linear_model import LinearRegression
lin_reg = LinearRegression()
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7ac2cec6a7505062, points: 2, grade: true, solution: false -->
```python
print("Linear regression parameters recovered by scikit-learn: intercept={0:.4f}, slope={1:.4f}"
      .format(lin_reg.intercept_, lin_reg.coef_[0]))
```

*Extend your model to include polynomial features up to degree 15 (using Scikit-Learn).  Use variable `lin_reg_poly` for your revised model.*

<!-- grade_id: cell-8ea9ea3be9d57ce5, points: 0, grade: false, solution: true -->
```python
degree = 15
bias = False
from sklearn.preprocessing import PolynomialFeatures
def compute_poly_features(degree, bias):
    # Return polynomial features of samples and class
    # YOUR CODE HERE
    raise NotImplementedError()
    return z_sample_poly, poly_features
z_sample_poly, poly_features = compute_poly_features(degree, bias)
# Train model
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d46184658a6b600b, points: 2, grade: true, solution: false -->
```python
check_var_defined('lin_reg_poly')
```

*Plot the data and the predictions of your models considered so far (linear and polynomial regression).  Call the variable used to store the predictions for your polynomial model `mu_pred_poly`.*

<!-- grade_id: cell-24f6bf2cfd9a7f71, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-55abac99185aa2df, points: 2, grade: true, solution: false -->
```python
check_var_defined('mu_pred_poly')
assert mu_pred_poly.shape == (len(z),)
```

*Comment on the accuracy of your models.*

<!-- grade_id: cell-c37fd48a9be1bd84, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

*Improve your polynomial model and use the improved model to make predictions. Call the variable used to store the polynomial model `ridge_reg_poly`. Call the variable used to store the predictions for your polynomial model `mu_pred_poly_improved`.*

<!-- grade_id: cell-249f8b263075ba25, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
check_var_defined('ridge_reg_poly')
check_var_defined('mu_pred_poly_improved')
assert mu_pred_poly_improved.shape == (len(z),), "Make sure the shape of your predictions is correct"
```

*Plot the predictions made with new model and all previous models considered.*

<!-- grade_id: cell-14a8a955921c3d9c, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

*Compute the RMS error between your predictions and the* data samples*.*

<!-- grade_id: cell-2ee5e0675e003f12, points: 0, grade: false, solution: true -->
```python
# Define a general function to compute the RMS error
def compute_rms(mu_1, mu_2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-946369d338039825, points: 1, grade: true, solution: false -->
```python
assert np.isclose(compute_rms(mu_pred_lin, mu_pred_lin), 0.0)
```

<!-- grade_id: cell-9c52009b6ad3fd17, points: 0, grade: false, solution: true -->
```python
# Compute the RMS error between the data and the predictions for each model.
# Use variables rms_sample_lin, rms_sample_poly and rms_sample_poly_improved.
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Print RMS values computed.
print("rms_sample_lin = {0:.4f}".format(rms_sample_lin))
print("rms_sample_poly = {0:.4f}".format(rms_sample_poly))
print("rms_sample_poly_improved = {0:.4f}".format(rms_sample_poly_improved))
```

<!-- grade_id: cell-09beb7cceccb40be, points: 1, grade: true, solution: false -->
```python
check_var_defined('rms_sample_lin')
```

<!-- grade_id: cell-6000d96ec4b0de43, points: 1, grade: true, solution: false -->
```python
check_var_defined('rms_sample_poly')
```

<!-- grade_id: cell-06d3a7838b6ad7e7, points: 1, grade: true, solution: false -->
```python
check_var_defined('rms_sample_poly_improved')
```

*Comment on what models you believe are best.*

<!-- grade_id: cell-d34e017b234e387c, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Using our cosmological concordance model we can predict the theoretical distance modulus vs redshift relationship using our understanding of the physics.

```python
from astroML.cosmology import Cosmology
cosmo = Cosmology()
mu_cosmo = np.array(list(map(cosmo.mu, z)))
```

*Plot the data, predictions made with all regression models, and the values predicted by the cosmological model.*

<!-- grade_id: cell-62cecf73f9f9a228, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

*Compute the RMS error between the predictions made by the cosmological model and each of the regression models, over the sample array `z`.*

<!-- grade_id: cell-939141aa04822e33, points: 0, grade: false, solution: true -->
```python
# Compute the RMS error between the data and the predictions for each model.
# Use variables rms_cosmo_lin, rms_cosmo_poly and rms_cosmo_poly_improved.
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Print RMS values computed.
print("rms_cosmo_lin = {0:.4f}".format(rms_cosmo_lin))
print("rms_cosmo_poly = {0:.4f}".format(rms_cosmo_poly))
print("rms_cosmo_poly_improved = {0:.4f}".format(rms_cosmo_poly_improved))
```

<!-- grade_id: cell-e8a9f757965b2069, points: 1, grade: true, solution: false -->
```python
check_var_defined('rms_cosmo_lin')
```

<!-- grade_id: cell-30ed009401260759, points: 1, grade: true, solution: false -->
```python
check_var_defined('rms_cosmo_poly')
```

<!-- grade_id: cell-e6ee872c00472aa2, points: 1, grade: true, solution: false -->
```python
check_var_defined('rms_cosmo_poly_improved')
```

*Comment on the RMS values computed and the implications for the accuracy of the different regression models considered.*

<!-- grade_id: cell-a993842383ec9778, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

---

## Part 2: Classification

In these exercises we will consider classification of [RR Lyrae](https://en.wikipedia.org/wiki/RR_Lyrae_variable) variable stars.  RR Lyrae variables are often used as standard candles to measure astronomical distances since their period of pulsation can be related to their absolute magnitude.

Observations of star magnitudes are made in each [SDSS filter band](http://skyserver.sdss.org/dr2/en/proj/advanced/color/sdssfilters.asp): u, g, r, i, z.

We will consider the space of astronomical "colours" to distinguish RR Lyraes from background stars.  Astronomical colours are simply differences in magnitudes between bands, e.g. u-g, g-r, r-i, i-z.  You can find further background [here](https://en.wikipedia.org/wiki/Color%E2%80%93color_diagram).

First, download the data.  (This may take some time on first execution.  Subsequently executions will read from cached data on your system.)

```python
# Load data
from astroML.datasets import fetch_rrlyrae_combined
X, y = fetch_rrlyrae_combined()
```

You can learn more about the format of the returned data [here](http://www.astroml.org/modules/generated/astroML.datasets.fetch_rrlyrae_combined.html).  In particular, note that the columns of `X` are u-g, g-r, r-i, i-z.

*Construct a Pandas DataFrame for the `X` data and a Series for the `y` data.  Call your Pandas objects `X_pd` and `y_pd` respectively.*

Be sure to give your colums the correct colour name, e.g. 'u-g'.

<!-- grade_id: cell-7250404d5b9e0c13, points: 0, grade: false, solution: true -->
```python
import pandas as pd
cols=['u-g', 'g-r', 'r-i', 'i-z']
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a913d8acabdfba5c, points: 1, grade: true, solution: false -->
```python
check_var_defined('X_pd')
X_pd.head()
```

<!-- grade_id: cell-0b438d5f8dcac8e9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d1392b89a707b35a, points: 1, grade: true, solution: false -->
```python
check_var_defined('y_pd')
y_pd.head()
```

*Combine your data and targets into a single Pandas DataFrame, labelling the target column 'target'.  Call the resulting Pandas DataFrame `X_pd_all`.*

<!-- grade_id: cell-f80ca9f2573d06fb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-694c65584675c6c6, points: 1, grade: true, solution: false -->
```python
check_var_defined('X_pd_all')
X_pd_all.head()
```

*Add a 'target description' column to your existing `X_pd_all` DataFrame, with fields 'Background' and 'RR Lyrae' to specify the target type.*

<!-- grade_id: cell-f94161f729fadf8c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cb2480e79d82c641, points: 1, grade: true, solution: false -->
```python
X_pd_all.head()
X_pd_all.tail()
```

*How many RR Lyrae variable stars are there in the dataset (i.e compute `n_rrlyrae`)?*

<!-- grade_id: cell-753a59a39e3df18c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c7fa425ec227dd04, points: 1, grade: true, solution: false -->
```python
check_var_defined('n_rrlyrae')
print("n_rrlyrae = {0}".format(n_rrlyrae))
```

*How many background stars are there in the dataset (i.e. compute `n_background`)?*

<!-- grade_id: cell-f902e74120d04b39, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-dd77ae406ebc1e36, points: 1, grade: true, solution: false -->
```python
check_var_defined('n_background')
print("n_background = {0}".format(n_background))
```

*Plot scatter plot pairs for all colour combinations using `seaborn`.  Colour the points by target type. Make sure the distribution plots are normalised to have an area of 1 under the curve for each of the classes.*

```python
%matplotlib inline
import seaborn as sns; sns.set()
```

<!-- grade_id: cell-6f8c0ce750628d0e, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Let's separate the data into training and test sets, keeping 25% of the data for testing.  

```python
from sklearn.model_selection import train_test_split
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.25, random_state=1)
```

First let's consider 1D classification for the zeroth colour, i.e. $u-g$. 

```python
ind = 0
col=cols[ind]
col
```

```python
X_train_1d = X_train[:, ind]
X_train_1d = X_train_1d.reshape(-1,1)
X_test_1d = X_test[:, ind]
X_test_1d = X_test_1d.reshape(-1,1)
```

To get some further intuition about the 1D classiciation problem consider a 1D plot of
class against colour.

```python
def plot_scatter():
    plt.figure(figsize=(10,5))
    plt.scatter(X_train_1d[y_train==1], y_train[y_train==1], c='m', marker='^', label='RR Lyrae')
    plt.scatter(X_train_1d[y_train==0], y_train[y_train==0], c='c', marker='v', label='Background')
    plt.xlabel('$' + col + '$')
    plt.ylabel('Probability of type RR Lyrae')
plot_scatter()    
plt.legend()
```

*Given the plot shown above, comment on how well you expect logistic regression to perform.*

<!-- grade_id: cell-1cad643fb7816037, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

*Where would you guess the decision bounary should lie?  Set the variable `decision_boundary_guess` to your guess.*

<!-- grade_id: cell-5eef717d4fab1828, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-00ef975a7880050f, points: 1, grade: true, solution: false -->
```python
check_var_defined('decision_boundary_guess')
print("decision_boundary_guess = {0:.4f}".format(decision_boundary_guess))
```

Use Scikit-Learn to perform logistic regression to classify the two classes for this 1D problem.

First, set the inverse regularation strength `C` such that regularisation is effecitvely not performed.

```python
C = 1e10
```

*Second, fit the model using Scikit-Learn. Use the variable `clf` for your classification model.*

<!-- grade_id: cell-f1790c24720c07d8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f6edcb4e5f610518, points: 1, grade: true, solution: false -->
```python
check_var_defined('clf')
```

*Compute the decision boundary of the logistic regression model fitted by Scikit-Learn.  User variable `decision_boundary_sklearn` for your result.*

(Ensure your result is a scalar and not an array of length 1.)

<!-- grade_id: cell-01cd8a3ebc69de43, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0ed39065189e2fae, points: 1, grade: true, solution: false -->
```python
check_var_defined('decision_boundary_sklearn')
assert not hasattr(decision_boundary_sklearn, "__len__")
print("decision_boundary_sklearn = {0:.4f}".format(decision_boundary_sklearn))
```

*Evaluate the probabilities predicted by your logistic regression model over the domain specified by the variable `X_1d_new`. Use variable `y_1d_proba` for your computed probabilities.*

<!-- grade_id: cell-b3f7cdf5d4698ad6, points: 0, grade: false, solution: true -->
```python
X_1d_new = np.linspace(0.3, 2.0, 1000).reshape(-1, 1)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bb76289d4e36fcb0, points: 1, grade: true, solution: false -->
```python
check_var_defined('y_1d_proba')
```

*Plot the probability of a star being of type RR Lyrae against the colour variable considered.  Also plot the probability of being a Background star.  Overlay these plots on the scatter plot of class types.  Also plot the decision boundary that you guessed previously and the one computed by Scikit-Learn.*

<!-- grade_id: cell-1c623b6df631aa69, points: 3, grade: true, solution: true -->
```python
plot_scatter()
# YOUR CODE HERE
raise NotImplementedError()
```

*From inspection of your plot, how would all objects in the training set be classified?*

<!-- grade_id: cell-dff437e03665f571, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

*Use your logistic regression model fitted by Scikit-Learn to predict the class of all objects in the test set. User variable `y_test_1d_pred` to specify your answer.*

<!-- grade_id: cell-bf444b0d8690c876, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d69905ed477cb96f, points: 1, grade: true, solution: false -->
```python
check_var_defined('y_test_1d_pred')
```

*How many objects are classified as of type RR Lyrae?  Use variable `n_rrlyrae_pred` to specify your answer.*

<!-- grade_id: cell-16f880b76044c462, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-357fb80562d278c5, points: 1, grade: true, solution: false -->
```python
check_var_defined('n_rrlyrae_pred')
assert n_rrlyrae_pred % 1 == 0 # check integer
print("n_rrlyrae_pred = {0}".format(n_rrlyrae_pred))
```

*How many objects are classified as of type Background?  Use variable `n_background_pred` to specify your answer.*

<!-- grade_id: cell-ba43adb513abebcd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5280ae78f8605c96, points: 1, grade: true, solution: false -->
```python
check_var_defined('n_background_pred')
assert n_background_pred % 1 == 0 # check integer
print("n_background_pred = {0}".format(n_background_pred))
```

Let's check the Scikit-Learn result by solving the logistic regression problem (without regularisation) manually.

Recall that the cost function for logistic regression is given by
$$
C(\theta) = -\frac{1}{m} \sum_{i=1}^m 
\left [ 
y^{(i)} \log(\hat{p}^{(i)})
+
(1 - y^{(i)}) \log(1 - \hat{p}^{(i)})
\right],
$$


where

$$\hat{p} = \sigma(\theta^\text{T} x) = \frac{1}{1+\exp{(-\theta^\text{T} x)}}. $$

Show analytically that the derivative of the cost function is given by
$$\begin{eqnarray}
\frac{\partial C}{\partial \theta} 
&=& 
\frac{1}{m} \sum_{i=1}^m 
\left[ \sigma\left(\theta^{\rm T} x^{(i)} \right) - y^{(i)} \right]
x^{(i)}\\
&=&
\frac{1}{m} 
X^{\rm T}
\left[ \sigma\left(X \theta \right) - y \right]
\end{eqnarray}$$

(use latex mathematics expressions).

*First, simplify the cost function terms $\log(\hat{p})$ and $\log(1-\hat{p})$ to express in terms linear in $\log\left({1+{\rm e}^{-\theta^{\rm T}x}}\right)$.*

(You may drop $i$ superscripts for notational brevity.)

<!-- grade_id: cell-5fc3a8343ec24488, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

*Next, substitute these terms into the cost function and simplify to also express the cost function in terms linear in $\log\left({1+{\rm e}^{-\theta^{\rm T}x}}\right)$.*

<!-- grade_id: cell-dcdb0de863dc8931, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

*Now compute the derivative of the cost function with respect to variable $\theta_j$, i.e. compute $\partial C / \partial \theta_j$.*

<!-- grade_id: cell-c386ea220c086ace, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

*Combine terms for all $\theta_j$ to give the overall derivative with respect to $\theta$, i.e. $\partial C / \partial \theta$.*

<!-- grade_id: cell-331a74ac412db42b, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

Using the analytically expression for the derivative of the cost function, we will solve the logistic regression problem by implementing a gradient descent algorithm.

*First, define the sigmoid function.*

<!-- grade_id: cell-e12fc0aa65b673b1, points: 0, grade: false, solution: true -->
```python
def sigmoid(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a5a50f4ec07d05fd, points: 1, grade: true, solution: false -->
```python
assert np.isclose(sigmoid(0), 0.5)
```

*Next, extend the training data to account for a bias term in your model. Use variable `X_train_1d_b` to specify your result.*

<!-- grade_id: cell-463d94ffced62fba, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c6f32c5137a9f302, points: 1, grade: true, solution: false -->
```python
check_var_defined('X_train_1d_b')
```

*Implement batch gradient descent to fit the parameters of your logistic regression model.  Consider `n_iterations = 40000` iterations and a learning rate of `alpha = 10.0`. Consider a starting point of $\theta_0 = (1, 1)$, i.e. `theta = np.array([[1], [1]])`. Use variable `theta` to specify your estimated parameters.*

<!-- grade_id: cell-aee503a999e27cf6, points: 0, grade: false, solution: true -->
```python
n_iterations = 40000
alpha = 10.0
theta = np.array([[1], [1]])
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-db0d3866ab6ed5aa, points: 4, grade: true, solution: false -->
```python
check_var_defined('theta')
print("theta[0] = {0:.4f}".format(theta[0][0]))
print("theta[1] = {0:.4f}".format(theta[1][0]))
```

*Compute the difference between the logistic regression model intercept computed by Scikit-Learn and manually.  Use variable `intercept_diff` for your result.*

<!-- grade_id: cell-b761dbdc7668fb7d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-eda44b051be24b4c, points: 1, grade: true, solution: false -->
```python
check_var_defined('intercept_diff')
print("intercept_diff = {0:.4E}".format(intercept_diff))
```

*Compute the difference between the logistic regression model* slope *(i.e. coefficient) computed by Scikit-Learn and manually.  Use variable `coeff_diff` for your result.*

<!-- grade_id: cell-17bd3970318abda0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-830185c3c51f3f91, points: 1, grade: true, solution: false -->
```python
check_var_defined('coeff_diff')
print("coeff_diff = {0:.4E}".format(coeff_diff))
```

You should find that the solution from your gradient descent algorithm is close (although not identical) to that recovered by Scikit-Learn. 

Both fitted logistic regression models, however, are not effective. The reason for this is because of class imbalance.  *Describe the class imbalance problem in your own words and how it manifests itself in the classification problem at hand.*

<!-- grade_id: cell-73126eae7fcd4d45, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

The class imbalance problem can be addressed by weighting the training data in a manner that is inversely proportional to their frequency.

*Repeat the fitting of your linear regression model but this time perform class weighting.  Use variable `clf_balanced` for your new model.*

See the `class_weight` argument of the Scikit-Learn [Logistic Regression](https://scikit-learn.org/stable/modules/generated/sklearn.linear_model.LogisticRegression.html) classifier for further details on how to perform class weighting.

<!-- grade_id: cell-043f89d606f8da67, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5dea5e84c6b3f90f, points: 0, grade: true, solution: false -->
```python
check_var_defined('clf_balanced')
```

*Compute the decision boundary of the logistic regression model fitted by Scikit-Learn when weighting classes.* 

(Ensure your result is a scalar and not an array of length 1.)

<!-- grade_id: cell-8789a822ce94928b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6d9e2c731edfff2f, points: 2, grade: true, solution: false -->
```python
check_var_defined('decision_boundary_sklearn_balanced')
assert not hasattr(decision_boundary_sklearn_balanced, "__len__")
print("decision_boundary_sklearn_balanced = {0:.4f}".format(decision_boundary_sklearn_balanced))
```

*Evaluate the probabilities prediced by your new logistic regression model over the domain specified by the variable `X_1d_new`. Use variable `y_1d_proba_balanced` for your computed probabilities.*

<!-- grade_id: cell-3db9585a121a321b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b5df6903e536bc3e, points: 1, grade: true, solution: false -->
```python
check_var_defined('y_1d_proba_balanced')
```

*For your new balanced model, plot the probability of a star being of type RR Lyrae against the colour variable considered.  Also plot the probability of being a Background star.  Overlay these plots on the scatter plot of class types.  Also plot the decision boundary that you guessed previously, the one computed by Scikit-Learn initially, and the one computed by Scikit-Learn for your new balanced model.*

<!-- grade_id: cell-9ce627c16d3996e6, points: 2, grade: true, solution: true -->
```python
plot_scatter()
# YOUR CODE HERE
raise NotImplementedError()
```

*Comment on the decision boundary of the balanced model compared to the unbalanced models fitted previously.*

<!-- grade_id: cell-45725a6e2c43d0f1, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

Now that we've built up good intuition surrounding the subtleties of the classification problem at hand in 1D, let's consider the 2D problem (we will keep to 2D for plotting convenience).

For the 2D case we consider the following colours.

```python
ind = 1
cols[:ind+1]
```

Consider the following training and test data for the 2D problem.

```python
X_train_2d = X_train[:, :ind+1]
X_train_2d = X_train_2d.reshape(-1,ind+1)
X_test_2d = X_test[:, :ind+1]
X_test_2d = X_test_2d.reshape(-1,ind+1)
```

*Train a logistic regression model for this 2D problem.  Use variable `clf_2d_logistic` for your classifier.*

<!-- grade_id: cell-807bce2068c8c513, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6d3421df624a8839, points: 0, grade: true, solution: false -->
```python
check_var_defined('clf_2d_logistic')
```

*Compute the precision and recall of your 2D logistic regression model. Use variables `precision_logistic` and `recall_logistic` for your results.*

<!-- grade_id: cell-5aa20025d9dcd3de, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-901a95a71017f77a, points: 1, grade: true, solution: false -->
```python
check_var_defined('precision_logistic')
print("precision_logistic = {0:.6f}".format(precision_logistic))
```

<!-- grade_id: cell-c829df4fb5b5646d, points: 1, grade: true, solution: false -->
```python
check_var_defined('recall_logistic')
print("recall_logistic = {0:.6f}".format(recall_logistic))
```

Consider the following meshgrid defining the u-g and g-r colour domain of interest.

```python
xlim = (0.7, 1.45)  # u-g
ylim = (-0.15, 0.4) # g-r
xx, yy = np.meshgrid(np.linspace(xlim[0], xlim[1], 100),
                     np.linspace(ylim[0], ylim[1], 100))
```

*Over the domain specified above plot the predicted classification probability.  Overlay on your plot the data instances, highlighting whether a RR Lyrae or background star, and the decision boundary.*

<!-- grade_id: cell-cb93ed5cd3864d37, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

*Now train an SVM classifier that can support a non-linear decision boundary on the same problem.*

<!-- grade_id: cell-c15b2b7d9de2fc9c, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

*Replicate for the SVM your plot above for the 2D logistic regression model.  Over the domain specified above plot the decision function score.  Overlay on your plot the data instances, highlighting whether a RR Lyrae or background star, and the decision boundary.*

<!-- grade_id: cell-3f1d0097c1a9a231, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

*Compute the precision and recall of your 2D SVM model. Use variables `precision_svm` and `recall_svm` for your results.*

<!-- grade_id: cell-4bb00fa2958c34b5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8a0f183fa2697432, points: 2, grade: true, solution: false -->
```python
check_var_defined('precision_svm')
print("precision_svm = {0:.6f}".format(precision_svm))
```

<!-- grade_id: cell-c88d680adad83470, points: 2, grade: true, solution: false -->
```python
check_var_defined('recall_svm')
print("recall_svm = {0:.6f}".format(recall_svm))
```

*Comment on the difference in decision boundary between your logistic regression and SVM models and how this impacts the effectiveness of the models.*

<!-- grade_id: cell-5429ab62dc4857e2, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE